<a href="https://colab.research.google.com/github/vidyaakkiraju/DS/blob/master/Vidya_Akkiraju_Python_Test_Level_0.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## 1(points 1)
Return the size of list.<br>
Example- list_size([1,2,3,4]) --> 4

<!-- grade_id: cell-4ea7ecd582660607, points: 0, grade: false, solution: true -->
```python
def list_size(l):
    # code here
    return len(l)

```

<!-- grade_id: cell-d03ede8261ade7b1, points: 1, grade: true, solution: false -->
```python
"""Testing code for last element"""
assert list_size([1,2,3,4,5 ]) == 5


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 2(points 1)
Return list of last five element of the given list.<br>
Example- def last_five(['a','b','c','d','e','f','g','h'])-->['d','e','f','g','h']

<!-- grade_id: cell-016db993367a12c3, points: 0, grade: false, solution: true -->
```python
def last_five(l):
    # code here
    return l[-5:]
```

<!-- grade_id: cell-50a0e61c4a912ec2, points: 1, grade: true, solution: false -->
```python
"""Testing code for last five elements"""
assert last_five([1,2,3,4,5,6,7]) == [3,4,5,6,7]


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 3(points 1)
Check whether given element exists in the list or not. If it exists the return True otherwise return False .<br>
Example - check_elem(['a','b','c','d','e','f','g','h'], 'a') -->True <br>
          check_elem(['a','b','c','d','e','f','g','h'], 'k') --> False

<!-- grade_id: cell-a69dc315a4b7af64, points: 0, grade: false, solution: true -->
```python
def check_elem(l,elem):
    # code here
    if elem in l:
      return True
    else:
      return False
```

<!-- grade_id: cell-8124565837ba565a, points: 0, grade: true, solution: false -->
```python
"""Testing code to check check_elem """
assert check_elem([1,2,3,4,5,6,7],10) == False


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 4(points 1)
Multiply the elements of a list and return the product.<br>
Example- mult_value([1,2,3,4,5])-->120

<!-- grade_id: cell-475ae2a35856e5a0, points: 0, grade: false, solution: true -->
```python
def mult_value(l):
    # code here
    k = 1
    for item in l:
      k *= item
    return k
```

<!-- grade_id: cell-53c0f4136b48a1e6, points: 1, grade: true, solution: false -->
```python
"""Testing code for mult_value """
assert mult_value([1,2,3,4]) == 24


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 5(points 1)
Return the number of occurences of given element in a list.<br>
Example- Element_rept([1,2,3,4,1,2,3,1,1,5] , 1)-->4<br>
Element_rept(['a','f','g','h','c','s','d','f','a','b','c','d'], c)--> 2

<!-- grade_id: cell-8b862b06e84772bf, points: 0, grade: false, solution: true -->
```python
def Element_rept(l,element):
    # code here
    '''
    count = 0
    for item in l:
      if item == element:
        count += 1
    return count '''
    return(l.count(element))
```

<!-- grade_id: cell-a702834de0ce86f7, points: 1, grade: true, solution: false -->
```python
"""Testing code to return the value after Multiplication """
assert Element_rept([1,4,1,2,3,1,5],10) == 0
assert Element_rept(['k','f','g','h','k','s','d','f','a','b','c','d'],'k') == 2

print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 6(points 1)
Return the number of keys in a dictionary.<br>
Example- dict_size({1:"akshay", 2:"vikash"})-->2

<!-- grade_id: cell-5f395a5882659f1c, points: 0, grade: false, solution: true -->
```python
def dict_size(d):
    # code here
    return len(d)
```

<!-- grade_id: cell-1ced6f5a600fae1a, points: 1, grade: true, solution: false -->
```python
"""Testing code to return the length"""

assert dict_size({1:1,2:2,3:3 }) == 3


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 7(points 1)
Return sorted list of keys of a dictionary.<br>
Examle- Sort_keyList({5:'IIT K',10:"IIT Jodhpur",3:"IIT Jaipur"})-->[3,5,10]

<!-- grade_id: cell-873ae424c0c727f7, points: 0, grade: false, solution: true -->
```python
def Sort_keyList(d):
    # code here
    '''
    li = []
    for key in d:
      li.append(key)
    return sorted(li)
    '''
    return sorted(d.keys())
```

<!-- grade_id: cell-d742cc724f24e259, points: 1, grade: true, solution: false -->
```python
"""Testing code to sort the dictionary"""
assert Sort_keyList({5:'IIT K',10:"IIT Jodhpur",3:"IIT Jaipur"}) == [3,5,10]


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 8(points 2)
Check wether given number is prime or not. If yes then retun 1, otherwise return 0.<br>
Example- is_prime(5) --> 1<br>
         is_prime(4) --> 0

<!-- grade_id: cell-3d0177977d340dd4, points: 0, grade: false, solution: true -->
```python
def is_prime(n):
    # code here
    if n == 2:
      return 1
    elif n ==1:
      return 0 
    else:
      for i in range(2,10):
        if (i != n) :
          if n%i == 0:
            return 0
            break
      return 1
```

<!-- grade_id: cell-ca1e99b17ed7b3c9, points: 2, grade: true, solution: false -->
```python
"""Testing code to check given number is prime or not """
assert is_prime(13) == 1


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 9(points 1)
Using the above fuction to return list of prime number from a to b (including both a and b).<br>
Example- primeNumberList(1,10)-->[2, 3, 5, 7]<br>
        

<!-- grade_id: cell-cc0a47cce1c28242, points: 0, grade: false, solution: true -->
```python
def primeNumberList(a,b):
    # code here
    lst = []
    for i in range(a,b):
      if is_prime(i) == 1:
        lst.append(i)
    return lst
      

```

<!-- grade_id: cell-7b15a712a22ba93b, points: 1, grade: true, solution: false -->
```python
"""Testing code to list of prime number """
assert primeNumberList(2,10) == [2, 3, 5, 7]


print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 10(points 2)
Return GCD of two number.<br>
Example- Gcd(2,10)-->2<br>
         Gcd(7,9)-->1

<!-- grade_id: cell-b594c4e142f9275c, points: 0, grade: false, solution: true -->
```python
def Gcd(a,b):
    # code here
    i = 1
    k = 1
    while (i <= a and i <= b):
      if (a%i == 0 and b%i == 0):
        k = i
      i = i + 1
    return k 
```

<!-- grade_id: cell-92b67715a65b4b0b, points: 2, grade: true, solution: false -->
```python
"""Testing code to return gcd"""
assert Gcd(2,10) == 2
 

print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

## 11(points 4)
Return a list of k-element combinations of elements in a list l

Example - comb([1, 2, 3, 4], 3) --> [[1, 2, 3], [1, 2, 4], [1, 3, 4], [2, 3, 4]]


<!-- grade_id: cell-a0a00f70e28b12f6, points: 0, grade: false, solution: true -->
```python
def comb(l, k):
  # code here
    from itertools import combinations_with_replacement
    kl = []
    m =  combinations_with_replacement(l,k)
    for i in list(m):
      kl.append((list(i)))
    return kl
```

<!-- grade_id: cell-44880bb9c9f2437e, points: 4, grade: true, solution: false -->
```python
"Testing code for comb"
assert [1, 2] in comb([1, 2, 3], 2)
assert [1, 3] in comb([1, 2, 3], 2)
assert [2, 3] in comb([1, 2, 3], 2)

print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

```python

```